# Data Structures and Processing

## Week2: Introduction to Python (Cont.)


## Note

Although we have covered the section on custom function or user defined function at the end of week 2's content, we would like to incorporate into practice right from beginning for two reasons:

* defining function becomes natural to you, and that you be able to extract out of your blocks a utility that could be used as many times as you please.

* we would like to run small tests against your functions to see if your utilities work as expected in several different scenarios.  This is called unit testing the functions, code or utility.

## Revisiting Personal Database

Recall that in the previous Jupyter notebook, we have asked you to make a small database, which you could use it for personal use.  We want to make it more functional and close to a real use case.

Let us describe the expectations and provide a framework.

* The database file `bibliox.txt` is stored on the disk.
* Each line of the file `biblio.txt` is a dictionary entry of a record.  Here is how a sample line in the file looks like.

    `{"id": "skiena2017data", "title": "The data science design manual", "authors": ["Steven Skiena"], "year": 2017, "publisher": "Springer"}`

We expect

* a function named `biblio_from_file` that could read the lines of the file `biblio.txt` and return a list of `dict` objects in Python.  The function should be able to take the name of the file as an argument.

* a function named `biblio_to_file` that takes a `list` of  `dict`s and writes its content to a given file name.  We use this function to ship the data to a file on computer.

* a function named `biblio_add_record` that takes a list as a first argument (the list of records) and a `dict` element (a new record) and adds the record to the list of others. This function should returs a list with additional record.

* a function named `biblio_record_presenet` that takes a list as a first argument (the list of records), and a `dict` element as a second argument and determines if a record is present in the list.

* a function named `biblio_query_record_by_author` that takes a list as a first argument (the list of records), and a `str` element as a second argument (the name of an author for an entry), and returns the list of all the matching records.

* a function named `biblio_remove_record` that takes a list as a first argument (the list of records), and a `dict` element as a second argument and removes it from the database. The function should return the biblio list without the record.


## Hints for a Solution to the Task on Personal Database

In the following sections, we are going to provide some hints that might be useful in writing down all the utilities desired for the task of a functional database.  You are suggested to think about them on your own and ignore the hints at first, but you are very welcomed to have a look at our suggestions, and provide your own solution.

## Hints: Reading the database file


### `biblio_from_file`

We have suggested to write down a function that could read the contents of a file and returns a list of records.

We suggest to break the task down into even smaller tasks or sub-tasks.  For example, we could use the `readlines()` method to read all the lines of the file into a `list` object.  This is easier.  If you look at an element of this list, you may notice that it is a `str` and contains a record together with some characters that we do not need, for example, the spaces or the braces `{`, `}`, or even the newline character `\n`. Therefore, at this point, we should concentrate on working on a helpful utility that could parse the line and return a `dict` element.  If we succeed in this, then we can use an iterative structure--the `for` loop, together with the newly defined utility to obtain the list of records as we need it.

Let us recall again how a line in the file looks like after it is read in.

        '{"id": "skiena2017data", "title": "The data science design manual", "authors": ["Steven Skiena"], "year": 2017, "publisher": "Springer"}\n'

Notice that the line is read as string but it contains different "strings" signified by the use of the doubl quotation marks `"`.

Let us write our observations and write utilities to act on them.

* The line contains the characters `{`, `}` and `\n`, which we do not need.  This means that we need to write down a function that would remove them.  The utility should return the cleaned string.

* After we have removed the undesired characters, the new string should look like the following

        '"id": "skiena2017data", "title": "The data science design manual", "authors": ["Steven Skiena"], "year": 2017, "publisher": "Springer"'
    
* In the cleaned line, we have a different parts that are meant to be the `key-value` pair for the `dict` object--our record.  We see that these are separated by a `,`.  Therefore, we need to find a way to break the string into several different parts and obtain a list of strings containing these `key-value` strings.  A helpful `str` method for such a purpose is the `split()`, which takes a string argument and splits that string at exactly that and returns a list. In our case, we want to split at `,`, therefore, we would use `split(",")` on the new line.  Our record after applying the utlity would look like the following.

        ['"id": "skiena2017data"',
         ' "title": "The data science design manual"',
         ' "authors": ["Steven Skiena"]',
         ' "year": 2017',
         ' "publisher": "Springer"']
  
  Notice that the second element in the above list is `' "title": "The data science design manual"'`. It has an extra white space at the beginning of the string.  We should make sure that our utility function takes care of it by removing the leading or the trailing spaces.  We do not know how many spaces there are, but they need to be removed.  A useful method for removing leading or trailing space is `strip()`. This method does not take any argument and removes the leading and the trailing spaces from a string, regardless of their number.

  If we have defined our function correctly, then the output should look lik the following.
  
        ['"id": "skiena2017data"',
         '"title": "The data science design manual"',
         '"authors": ["Steven Skiena"]',
         '"year": 2017',
         '"publisher": "Springer"']
         
* Now need to actually take the members of the list and write down a utility that would return a dictionary.  We could divide this task into subtasks, where we first concentrate on a member of a list which should provide the `key-value` pair.  For example, let us concentrate on the string `'"title": "The data science design manual"`.  There might be several ways to go forward from here.  But one interesting way is to use our previosly defined utility which splitted the given string at `,`, removed the leading and trailing spaces, and returned a list.  Notice that we need the same utility but with a small change that the character at which we want to split now is not `,`, but it is `:`.  This takes us one step back, and motivates us to add the splitting charachter a part of our input parameters of our utility, because it would then provide us more control.

  If the utility is redefined, or another utility is defined instead, we should have as our result the following list

        ['title', 'The data science design manual']
        
   We can the first element (at index 0) as the key and the second element (at index 1) as the corresponding value to build a dictionary or add it to a dictionary.
  


```python
import json

def biblio_from_file(filename: str):

    records = []
    try:
        with open(filename, "r", encoding="utf-8") as file:
            for line in file:
                line = line.strip()
                if line:
                    record = json.loads(line)
                    records.append(record)
    except FileNotFoundError:
        print(f"Error: The file {filename} was not found.")
    except json.JSONDecodeError as e:
        print(f"Error in the format JSON : {e}")

    return records

```

```python
assert type(biblio_from_file("biblio.txt")) == list, "The function does not return a list, as expected."
assert type(biblio_from_file("biblio.txt")[0]) == dict, "The list does not contain dictionaries."
```

Output:
```
Error: The file biblio.txt was not found.
Error: The file biblio.txt was not found.
IndexError                                Traceback (most recent call last)
Cell In[3], line 2
      1 assert type(biblio_from_file("biblio.txt")) == list, "The function does not return a list, as expected."
----> 2 assert type(biblio_from_file("biblio.txt")[0]) == dict, "The list does not contain dictionaries."

IndexError: list index out of range
```

```python
biblio = biblio_from_file("biblio.txt")
print(biblio)

```

 By this point, you should be able to put all the utilities, defined above, together with iterative loops to build the function `biblio_from_file`.

### `biblio_to_file`

This function might be easier to write as compared to the previous `biblio_from_file`.

The situaiton is that we have our record `biblio`, which is a list of `dict` objects – the records. We should be able to use the `for` statement to write each element of the list `biblio` as a line in the output file.  We leave the details for you to think and provide the solution.

<!-- grade_id: cell-6b70252b4f8e2376, points: 0, grade: false, solution: true -->
```python
import json

def biblio_to_file(biblio: list, filename: str):

    try:
        with open(filename, "w", encoding="utf-8") as file:
            for record in biblio:
                json.dump(record, file)
                file.write("\n")
    except Exception as e:
        print(f"Error to save the file: {e}")
biblio_to_file(biblio_data, "biblio.txt")
print("File saved properly")


```

Output:
```
File saved properly
```

### `biblio_add_record`

This function is also easier to write.

The situation is that we have a list and we want to add an element to it. A better version would be to see if the record is also available in `biblio`, before actually adding it.  There is no need to add a record if it is already a part of the database.  Therefore, before defining this function, we should defined the following which checks if a record is already present.

<!-- grade_id: cell-d075412b8b44fb76, points: 0, grade: false, solution: true -->
```python
def biblio_add_record(biblio: list, record: dict):

    if record not in biblio:
        biblio.append(record)
    return biblio
```

<!-- grade_id: cell-65ed98d7df10f024, points: 5, grade: true, solution: false -->
```python
biblio_data = biblio_from_file("biblio.txt")
biblio_data = biblio_add_record(biblio_data, xrecord)

assert xrecord in biblio_data, "Error: The record was not added"
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[23], line 2
      1 biblio_data = biblio_from_file("biblio.txt")
----> 2 biblio_data = biblio_add_record(biblio_data, xrecord)
      4 assert xrecord in biblio_data, "Error: The record was not added"

NameError: name 'xrecord' is not defined
```

### `biblio_record_present`

This function is one the simplest.  But we would like to have it for completion of our system.

The situation is that we have a list and we want to check an object is in the list.  We use the infix operator `in` which returns one of the boolean values: `True` or `False`.

<!-- grade_id: cell-d06e3cd3cc3314f8, points: 0, grade: false, solution: true -->
```python
def biblio_record_present(biblio: list, record: dict) -> bool:

    return record in biblio
```

```python
print("Data base content", biblio_data)

```

### `biblio_query_record_by_author`

This function filters for matching records and returns a list.

To define this function, we basically rely on the `for` statement which would take each record (a `dict` object), let us call it `record`, and gets the value of authors by `record["author"]`, which is a list of authors for a particular bibliographic entry.  We can then use our given string together with the `in` operator to determine the corresponding boolean value.  `True` would mean that the the `record` matches and we should append it to a temporary list that we are going to return.  We append using the `list` method `append()`. If `False` is returned, then it means that we ignore it.

<!-- grade_id: cell-b3cb5e8e421ba323, points: 0, grade: false, solution: true -->
```python
def biblio_query_record_by_author(biblio: list, author: str) -> list:

    return [record for record in biblio if "authors" in record and author in record["authors"]]

```

### `biblio_query_record_by_key` (optional)

We have not asked for this function, but it is worth thinking that we might want to retrieve list of records on the basis of other criteria than just the `author`.  One way is to define function for all possible keys that we could include, or the second option is to define a generate function that takes `key-value` pair as arguments and returns all the matched records.

After having the `biblio_query_record_by_author` defined, it is not hard to define the abstracted version.  We leave it as a small exercise to think how to best do it.

```python
def biblio_query_record_by_key(biblio: list, key: str, value) -> list:

    return [record for record in biblio if key in record and record[key] == value]

```

### `biblio_remove_record`

The simplest thinking here is that we have a list of records and we want to remove the records if it is present.  Our function `biblio_record_present`, already defined, might be helpful to check if a record is present.

Another way could be to use the `list` method `remove()` that takes the object as an argument that we want to remove.  There is a need to be careful here. If the object that we are trying to remove is not present in the first place, then the method `remove()` returns `TypeError`.  At this point, we can either our utility function or use the `try...except` to avoid the error.

<!-- grade_id: cell-733a1907ce49b3c5, points: 0, grade: false, solution: true -->
```python
def biblio_remove_record(biblio: list, record: dict) -> list:

    if record in biblio:
        biblio.remove(record)
    return biblio

```

<!-- grade_id: cell-b1db669c4f9e8d72, points: 5, grade: true, solution: false -->
```python

biblio_data = biblio_from_file("biblio.txt")


biblio_data = biblio_remove_record(biblio_data, xrecord)


assert xrecord not in biblio_data, "Error: The record was not deleted successfully"


biblio_to_file(biblio_data, "biblio.txt")

print(" Record deleted successfully.")

```

# Additional Task
 You are now familiar with using `.biblio` and are comfortable managing records within it. But are you ready to solve harder tasks?

## Create a student management system for your classroom database.
## Follow the steps below:

1) Create a Student Model. Define a student model with the following attributes:

    * matriculation_number
    * first_name
    * last_name
    * country_of_origin
    * enrollment_year
    * major

2) Create at least 3 student records and store them in your `.biblio` file (representing the classroom).

3) Convert the `.biblio` file into a string format and print the result.

4) Remove 2 students from the classroom and update the`.biblio` file.

5) Convert the updated `.biblio` file to a string and print it again to show the changes.

### Hint: Use the previously defined functions to complete this task.

For creating the student model, you can work off a snippet of the function you would need.


```python
def create_student(matriculation_number, first_name, last_name, country_of_origin, enrollment_year, major):
    return {
        "matriculation_number": matriculation_number,
        "first_name": first_name,
        "last_name": last_name,
        "country_of_origin": country_of_origin,
        "enrollment_year": enrollment_year,
        "major": major
    }
```

```python
student1 = create_student("30001234", "Henry", "Cavill", "USA", 2001, "IBA")
student2 = create_student("30005678", "Gal", "Gadot", "Israel", 2002, "CS")
student3 = create_student("30009012", "Jason", "Momoa", "Hawaii", 2003, "ME")

biblio_data = biblio_from_file("biblio.txt")
biblio_data = biblio_add_record(biblio_data, student1)
biblio_data = biblio_add_record(biblio_data, student2)
biblio_data = biblio_add_record(biblio_data, student3)
biblio_to_file(biblio_data, "biblio.txt")
```

```python
with open("biblio.txt", "r", encoding="utf-8") as file:
    biblio_string = file.read()
print(biblio_string)
```

Output:
```
{"matriculation_number": "30009012", "first_name": "Jason", "last_name": "Momoa", "country_of_origin": "Hawaii", "enrollment_year": 2003, "major": "ME"}
{"matriculation_number": "30001234", "first_name": "Henry", "last_name": "Cavill", "country_of_origin": "USA", "enrollment_year": 2001, "major": "IBA"}
{"matriculation_number": "30005678", "first_name": "Gal", "last_name": "Gadot", "country_of_origin": "Israel", "enrollment_year": 2002, "major": "CS"}
```

```python
biblio_data = biblio_remove_record(biblio_data, student1)
biblio_data = biblio_remove_record(biblio_data, student2)
biblio_to_file(biblio_data, "biblio.txt")
```

```python
with open("biblio.txt", "r", encoding="utf-8") as file:
    biblio_string = file.read()
print(biblio_string)
```

Output:
```
{"matriculation_number": "30009012", "first_name": "Jason", "last_name": "Momoa", "country_of_origin": "Hawaii", "enrollment_year": 2003, "major": "ME"}
```

```python

biblio_to_file(biblio_data, "biblio.txt")


with open("biblio.txt", "r", encoding="utf-8") as file:
    biblio_data = [json.loads(line.strip()) for line in file]
biblio_string = str(biblio_data)

print("String printed dictionary",biblio_string)
```

Output:
```
String printed dictionary [{'matriculation_number': '30009012', 'first_name': 'Jason', 'last_name': 'Momoa', 'country_of_origin': 'Hawaii', 'enrollment_year': 2003, 'major': 'ME'}]
```